---
title: 'Lesson 2.2: Capture output of shell commands'
description: ' Capture output of shell commands using subprocess.Popen '
---

## Example coding exercise

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

This is an example exercise.

`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---

## Capture output of shell commands

```yaml
type: VideoExercise
key: 4709e638fc
xp: 50
```

`@projector_key`
0e77cf66e9b9d1c38fee414ce04f1431
